% dis+11_10_add=off:afr=on:afp=100000:afq=1.2:amm=off:anc=none:cond=on:gs=on:gsem=on:inw=on:ile=on:irw=on:lma=on:nm=64:newcnf=on:nwc=10:sas=z3:sac=on:sp=reverse_arity_2 on nia_cubes_simple_v2
% Refutation found. Thanks to Tanya!
% SZS status Unsatisfiable for nia_cubes_simple_v2
% SZS output start Proof for nia_cubes_simple_v2
1. ? [X1 : $int,X3 : $int,X2 : $int,X0 : $int] : (0 = $difference($difference(X2,$product(6,X3)),6) & 0 = $sum($difference($difference($sum($difference($product(X1,X1),$product($product(9,X3),X0)),$product(9,X3)),$product(5,X1)),$product(18,X0)),4) & 0 = $sum($difference($difference($sum($product(X3,X1),$product(2,X3)),X1),$product(3,X0)),1) & 0 = $sum($difference($sum($product($product(3,X3),X3),$product(3,X3)),X1),1) & 0 = $difference($difference($sum(X2,6),$product(6,$sum(X3,1))),6) & 0 = $sum($difference($difference($sum($difference($product($sum(X1,X2),$sum(X1,X2)),$product($product(9,$sum(X3,1)),$sum(X0,X1))),$product(9,$sum(X3,1))),$product(5,$sum(X1,X2))),$product(18,$sum(X0,X1))),4) & (~0 = $sum($difference($difference($sum($product($sum(X3,1),$sum(X1,X2)),$product(2,$sum(X3,1))),$sum(X1,X2)),$product(3,$sum(X0,X1))),1) | ~0 = $sum($difference($sum($product($product(3,$sum(X3,1)),$sum(X3,1)),$product(3,$sum(X3,1))),$sum(X1,X2)),1))) [input]
2. ? [X1 : $int,X3 : $int,X2 : $int,X0 : $int] : (0 = $sum($difference(X2,$product(6,X3)),$uminus(6)) & 0 = $sum($sum($difference($sum($difference($product(X1,X1),$product($product(9,X3),X0)),$product(9,X3)),$product(5,X1)),$uminus($product(18,X0))),4) & 0 = $sum($sum($difference($sum($product(X3,X1),$product(2,X3)),X1),$uminus($product(3,X0))),1) & 0 = $sum($sum($sum($product($product(3,X3),X3),$product(3,X3)),$uminus(X1)),1) & 0 = $sum($difference($sum(X2,6),$product(6,$sum(X3,1))),$uminus(6)) & 0 = $sum($sum($difference($sum($difference($product($sum(X1,X2),$sum(X1,X2)),$product($product(9,$sum(X3,1)),$sum(X0,X1))),$product(9,$sum(X3,1))),$product(5,$sum(X1,X2))),$uminus($product(18,$sum(X0,X1)))),4) & (~0 = $sum($sum($difference($sum($product($sum(X3,1),$sum(X1,X2)),$product(2,$sum(X3,1))),$sum(X1,X2)),$uminus($product(3,$sum(X0,X1)))),1) | ~0 = $sum($sum($sum($product($product(3,$sum(X3,1)),$sum(X3,1)),$product(3,$sum(X3,1))),$uminus($sum(X1,X2))),1))) [theory normalization 1]
3. $sum(X0,X1) = $sum(X1,X0) [theory axiom 127]
14. $product(X0,X1) = $product(X1,X0) [theory axiom 127]
21. ? [X1 : $int,X3 : $int,X2 : $int,X0 : $int] : (0 = $sum($sum(X2,$uminus($product(6,X3))),$uminus(6)) & 0 = $sum($sum($sum($sum($difference($product(X1,X1),$product($product(9,X3),X0)),$product(9,X3)),$uminus($product(5,X1))),$uminus($product(18,X0))),4) & 0 = $sum($sum($sum($sum($product(X3,X1),$product(2,X3)),$uminus(X1)),$uminus($product(3,X0))),1) & 0 = $sum($sum($sum($product($product(3,X3),X3),$product(3,X3)),$uminus(X1)),1) & 0 = $sum($sum($sum(X2,6),$uminus($product(6,$sum(X3,1)))),$uminus(6)) & 0 = $sum($sum($sum($sum($difference($product($sum(X1,X2),$sum(X1,X2)),$product($product(9,$sum(X3,1)),$sum(X0,X1))),$product(9,$sum(X3,1))),$uminus($product(5,$sum(X1,X2)))),$uminus($product(18,$sum(X0,X1)))),4) & (~0 = $sum($sum($sum($sum($product($sum(X3,1),$sum(X1,X2)),$product(2,$sum(X3,1))),$uminus($sum(X1,X2))),$uminus($product(3,$sum(X0,X1)))),1) | ~0 = $sum($sum($sum($product($product(3,$sum(X3,1)),$sum(X3,1)),$product(3,$sum(X3,1))),$uminus($sum(X1,X2))),1))) [theory normalization 2]
22. ? [X0 : $int,X1 : $int,X2 : $int,X3 : $int] : (0 = $sum($sum(X2,$uminus($product(6,X1))),$uminus(6)) & 0 = $sum($sum($sum($sum($difference($product(X0,X0),$product($product(9,X1),X3)),$product(9,X1)),$uminus($product(5,X0))),$uminus($product(18,X3))),4) & 0 = $sum($sum($sum($sum($product(X1,X0),$product(2,X1)),$uminus(X0)),$uminus($product(3,X3))),1) & 0 = $sum($sum($sum($product($product(3,X1),X1),$product(3,X1)),$uminus(X0)),1) & 0 = $sum($sum($sum(X2,6),$uminus($product(6,$sum(X1,1)))),$uminus(6)) & 0 = $sum($sum($sum($sum($difference($product($sum(X0,X2),$sum(X0,X2)),$product($product(9,$sum(X1,1)),$sum(X3,X0))),$product(9,$sum(X1,1))),$uminus($product(5,$sum(X0,X2)))),$uminus($product(18,$sum(X3,X0)))),4) & (~0 = $sum($sum($sum($sum($product($sum(X1,1),$sum(X0,X2)),$product(2,$sum(X1,1))),$uminus($sum(X0,X2))),$uminus($product(3,$sum(X3,X0)))),1) | ~0 = $sum($sum($sum($product($product(3,$sum(X1,1)),$sum(X1,1)),$product(3,$sum(X1,1))),$uminus($sum(X0,X2))),1))) [rectify 21]
23. ? [X0 : $int,X1 : $int,X2 : $int,X3 : $int] : (0 = $sum($sum(X2,$uminus($product(6,X1))),$uminus(6)) & 0 = $sum($sum($sum($sum($difference($product(X0,X0),$product($product(9,X1),X3)),$product(9,X1)),$uminus($product(5,X0))),$uminus($product(18,X3))),4) & 0 = $sum($sum($sum($sum($product(X1,X0),$product(2,X1)),$uminus(X0)),$uminus($product(3,X3))),1) & 0 = $sum($sum($sum($product($product(3,X1),X1),$product(3,X1)),$uminus(X0)),1) & 0 = $sum($sum($sum(X2,6),$uminus($product(6,$sum(X1,1)))),$uminus(6)) & 0 = $sum($sum($sum($sum($difference($product($sum(X0,X2),$sum(X0,X2)),$product($product(9,$sum(X1,1)),$sum(X3,X0))),$product(9,$sum(X1,1))),$uminus($product(5,$sum(X0,X2)))),$uminus($product(18,$sum(X3,X0)))),4) & (0 != $sum($sum($sum($sum($product($sum(X1,1),$sum(X0,X2)),$product(2,$sum(X1,1))),$uminus($sum(X0,X2))),$uminus($product(3,$sum(X3,X0)))),1) | 0 != $sum($sum($sum($product($product(3,$sum(X1,1)),$sum(X1,1)),$product(3,$sum(X1,1))),$uminus($sum(X0,X2))),1))) [flattening 22]
24. 0 != $sum($sum($sum($product($product(3,$sum(sK1,1)),$sum(sK1,1)),$product(3,$sum(sK1,1))),$uminus($sum(sK0,sK2))),1) | 0 != $sum($sum($sum($sum($product($sum(sK1,1),$sum(sK0,sK2)),$product(2,$sum(sK1,1))),$uminus($sum(sK0,sK2))),$uminus($product(3,$sum(sK3,sK0)))),1) [cnf transformation 23]
25. 0 = $sum($sum($sum($sum($difference($product($sum(sK0,sK2),$sum(sK0,sK2)),$product($product(9,$sum(sK1,1)),$sum(sK3,sK0))),$product(9,$sum(sK1,1))),$uminus($product(5,$sum(sK0,sK2)))),$uminus($product(18,$sum(sK3,sK0)))),4) [cnf transformation 23]
26. 0 = $sum($sum($sum(sK2,6),$uminus($product(6,$sum(sK1,1)))),$uminus(6)) [cnf transformation 23]
27. 0 = $sum($sum($sum($product($product(3,sK1),sK1),$product(3,sK1)),$uminus(sK0)),1) [cnf transformation 23]
28. 0 = $sum($sum($sum($sum($product(sK1,sK0),$product(2,sK1)),$uminus(sK0)),$uminus($product(3,sK3))),1) [cnf transformation 23]
29. 0 = $sum($sum($sum($sum($difference($product(sK0,sK0),$product($product(9,sK1),sK3)),$product(9,sK1)),$uminus($product(5,sK0))),$uminus($product(18,sK3))),4) [cnf transformation 23]
30. 0 = $sum($sum(sK2,$uminus($product(6,sK1))),$uminus(6)) [cnf transformation 23]
32. 0 = $sum(sK2,$uminus($product(6,$sum(sK1,1)))) [evaluation 26]
33. 0 = $sum($sum(sK2,$uminus($product(6,sK1))),-6) [evaluation 30]
408. 0 = $sum(sK2,$uminus($product(6,$sum(1,sK1)))) [forward demodulation 32,3]
455. 1 <=> 0 = $sum(sK2,$uminus($product(6,$sum(1,sK1)))) [avatar definition]
458. 1 [avatar split clause 408,455]
552. 2 <=> 0 = $sum($sum(sK2,$uminus($product(6,sK1))),-6) [avatar definition]
555. 2 [avatar split clause 33,552]
977. 0 = $sum($sum($uminus(sK0),$sum($product($product(3,sK1),sK1),$product(3,sK1))),1) [forward demodulation 27,3]
978. 0 = $sum($sum($uminus(sK0),$sum($product(3,sK1),$product($product(3,sK1),sK1))),1) [forward demodulation 977,3]
979. 0 = $sum($sum($uminus(sK0),$sum($product(3,sK1),$product(sK1,$product(3,sK1)))),1) [forward demodulation 978,14]
1028. 3 <=> 0 = $sum($sum($uminus(sK0),$sum($product(3,sK1),$product(sK1,$product(3,sK1)))),1) [avatar definition]
1031. 3 [avatar split clause 979,1028]
1168. 0 = $sum($sum($uminus($product(3,sK3)),$sum($sum($product(sK1,sK0),$product(2,sK1)),$uminus(sK0))),1) [forward demodulation 28,3]
1169. 0 = $sum($sum($uminus($product(3,sK3)),$sum($uminus(sK0),$sum($product(sK1,sK0),$product(2,sK1)))),1) [forward demodulation 1168,3]
1170. 0 = $sum($sum($uminus($product(3,sK3)),$sum($uminus(sK0),$sum($product(2,sK1),$product(sK1,sK0)))),1) [forward demodulation 1169,3]
1171. 0 = $sum($sum($uminus($product(3,sK3)),$sum($uminus(sK0),$sum($product(2,sK1),$product(sK0,sK1)))),1) [forward demodulation 1170,14]
1233. 4 <=> 0 = $sum($sum($uminus($product(3,sK3)),$sum($uminus(sK0),$sum($product(2,sK1),$product(sK0,sK1)))),1) [avatar definition]
1236. 4 [avatar split clause 1171,1233]
1360. 0 = $sum($sum($uminus($product(18,sK3)),$sum($sum($difference($product(sK0,sK0),$product($product(9,sK1),sK3)),$product(9,sK1)),$uminus($product(5,sK0)))),4) [forward demodulation 29,3]
1361. 0 = $sum($sum($uminus($product(18,sK3)),$sum($uminus($product(5,sK0)),$sum($difference($product(sK0,sK0),$product($product(9,sK1),sK3)),$product(9,sK1)))),4) [forward demodulation 1360,3]
1362. 0 = $sum($sum($uminus($product(18,sK3)),$sum($uminus($product(5,sK0)),$sum($product(9,sK1),$difference($product(sK0,sK0),$product($product(9,sK1),sK3))))),4) [forward demodulation 1361,3]
1363. 0 = $sum($sum($uminus($product(18,sK3)),$sum($uminus($product(5,sK0)),$sum($product(9,sK1),$difference($product(sK0,sK0),$product(sK3,$product(9,sK1)))))),4) [forward demodulation 1362,14]
1422. 5 <=> 0 = $sum($sum($uminus($product(18,sK3)),$sum($uminus($product(5,sK0)),$sum($product(9,sK1),$difference($product(sK0,sK0),$product(sK3,$product(9,sK1)))))),4) [avatar definition]
1425. 5 [avatar split clause 1363,1422]
1537. 0 = $sum($sum($uminus($product(18,$sum(sK3,sK0))),$sum($sum($difference($product($sum(sK0,sK2),$sum(sK0,sK2)),$product($product(9,$sum(sK1,1)),$sum(sK3,sK0))),$product(9,$sum(sK1,1))),$uminus($product(5,$sum(sK0,sK2))))),4) [forward demodulation 25,3]
1538. 0 = $sum($sum($uminus($product(18,$sum(sK3,sK0))),$sum($uminus($product(5,$sum(sK0,sK2))),$sum($difference($product($sum(sK0,sK2),$sum(sK0,sK2)),$product($product(9,$sum(sK1,1)),$sum(sK3,sK0))),$product(9,$sum(sK1,1))))),4) [forward demodulation 1537,3]
1539. 0 = $sum($sum($uminus($product(18,$sum(sK3,sK0))),$sum($uminus($product(5,$sum(sK0,sK2))),$sum($product(9,$sum(sK1,1)),$difference($product($sum(sK0,sK2),$sum(sK0,sK2)),$product($product(9,$sum(sK1,1)),$sum(sK3,sK0)))))),4) [forward demodulation 1538,3]
1540. 0 = $sum($sum($uminus($product(18,$sum(sK3,sK0))),$sum($uminus($product(5,$sum(sK0,sK2))),$sum($product(9,$sum(sK1,1)),$difference($product($sum(sK0,sK2),$sum(sK0,sK2)),$product($sum(sK3,sK0),$product(9,$sum(sK1,1))))))),4) [forward demodulation 1539,14]
1541. 0 = $sum($sum($uminus($product(18,$sum(sK3,sK0))),$sum($uminus($product(5,$sum(sK0,sK2))),$sum($product(9,$sum(1,sK1)),$difference($product($sum(sK0,sK2),$sum(sK0,sK2)),$product($sum(sK3,sK0),$product(9,$sum(1,sK1))))))),4) [forward demodulation 1540,3]
1542. 0 = $sum($sum($uminus($product(18,$sum(sK0,sK3))),$sum($uminus($product(5,$sum(sK0,sK2))),$sum($product(9,$sum(1,sK1)),$difference($product($sum(sK0,sK2),$sum(sK0,sK2)),$product($sum(sK0,sK3),$product(9,$sum(1,sK1))))))),4) [forward demodulation 1541,3]
1583. 6 <=> 0 = $sum($sum($uminus($product(18,$sum(sK0,sK3))),$sum($uminus($product(5,$sum(sK0,sK2))),$sum($product(9,$sum(1,sK1)),$difference($product($sum(sK0,sK2),$sum(sK0,sK2)),$product($sum(sK0,sK3),$product(9,$sum(1,sK1))))))),4) [avatar definition]
1586. 6 [avatar split clause 1542,1583]
1695. 0 != $sum(1,$sum($sum($product($product(3,$sum(sK1,1)),$sum(sK1,1)),$product(3,$sum(sK1,1))),$uminus($sum(sK0,sK2)))) | 0 != $sum($sum($sum($sum($product($sum(sK1,1),$sum(sK0,sK2)),$product(2,$sum(sK1,1))),$uminus($sum(sK0,sK2))),$uminus($product(3,$sum(sK3,sK0)))),1) [forward demodulation 24,3]
1696. 0 != $sum(1,$sum($uminus($sum(sK0,sK2)),$sum($product($product(3,$sum(sK1,1)),$sum(sK1,1)),$product(3,$sum(sK1,1))))) | 0 != $sum($sum($sum($sum($product($sum(sK1,1),$sum(sK0,sK2)),$product(2,$sum(sK1,1))),$uminus($sum(sK0,sK2))),$uminus($product(3,$sum(sK3,sK0)))),1) [forward demodulation 1695,3]
1697. 0 != $sum(1,$sum($uminus($sum(sK0,sK2)),$sum($product(3,$sum(sK1,1)),$product($product(3,$sum(sK1,1)),$sum(sK1,1))))) | 0 != $sum($sum($sum($sum($product($sum(sK1,1),$sum(sK0,sK2)),$product(2,$sum(sK1,1))),$uminus($sum(sK0,sK2))),$uminus($product(3,$sum(sK3,sK0)))),1) [forward demodulation 1696,3]
1698. 0 != $sum(1,$sum($uminus($sum(sK0,sK2)),$sum($product(3,$sum(sK1,1)),$product($sum(sK1,1),$product(3,$sum(sK1,1)))))) | 0 != $sum($sum($sum($sum($product($sum(sK1,1),$sum(sK0,sK2)),$product(2,$sum(sK1,1))),$uminus($sum(sK0,sK2))),$uminus($product(3,$sum(sK3,sK0)))),1) [forward demodulation 1697,14]
1699. 0 != $sum(1,$sum($uminus($sum(sK0,sK2)),$sum($product(3,$sum(1,sK1)),$product($sum(1,sK1),$product(3,$sum(1,sK1)))))) | 0 != $sum($sum($sum($sum($product($sum(sK1,1),$sum(sK0,sK2)),$product(2,$sum(sK1,1))),$uminus($sum(sK0,sK2))),$uminus($product(3,$sum(sK3,sK0)))),1) [forward demodulation 1698,3]
1700. 0 != $sum($sum($uminus($product(3,$sum(sK3,sK0))),$sum($sum($product($sum(sK1,1),$sum(sK0,sK2)),$product(2,$sum(sK1,1))),$uminus($sum(sK0,sK2)))),1) | 0 != $sum(1,$sum($uminus($sum(sK0,sK2)),$sum($product(3,$sum(1,sK1)),$product($sum(1,sK1),$product(3,$sum(1,sK1)))))) [forward demodulation 1699,3]
1701. 0 != $sum($sum($uminus($product(3,$sum(sK3,sK0))),$sum($uminus($sum(sK0,sK2)),$sum($product($sum(sK1,1),$sum(sK0,sK2)),$product(2,$sum(sK1,1))))),1) | 0 != $sum(1,$sum($uminus($sum(sK0,sK2)),$sum($product(3,$sum(1,sK1)),$product($sum(1,sK1),$product(3,$sum(1,sK1)))))) [forward demodulation 1700,3]
1702. 0 != $sum($sum($uminus($product(3,$sum(sK3,sK0))),$sum($uminus($sum(sK0,sK2)),$sum($product(2,$sum(sK1,1)),$product($sum(sK1,1),$sum(sK0,sK2))))),1) | 0 != $sum(1,$sum($uminus($sum(sK0,sK2)),$sum($product(3,$sum(1,sK1)),$product($sum(1,sK1),$product(3,$sum(1,sK1)))))) [forward demodulation 1701,3]
1703. 0 != $sum($sum($uminus($product(3,$sum(sK3,sK0))),$sum($uminus($sum(sK0,sK2)),$sum($product(2,$sum(sK1,1)),$product($sum(sK0,sK2),$sum(sK1,1))))),1) | 0 != $sum(1,$sum($uminus($sum(sK0,sK2)),$sum($product(3,$sum(1,sK1)),$product($sum(1,sK1),$product(3,$sum(1,sK1)))))) [forward demodulation 1702,14]
1704. 0 != $sum($sum($uminus($product(3,$sum(sK3,sK0))),$sum($uminus($sum(sK0,sK2)),$sum($product(2,$sum(1,sK1)),$product($sum(sK0,sK2),$sum(1,sK1))))),1) | 0 != $sum(1,$sum($uminus($sum(sK0,sK2)),$sum($product(3,$sum(1,sK1)),$product($sum(1,sK1),$product(3,$sum(1,sK1)))))) [forward demodulation 1703,3]
1705. 0 != $sum($sum($uminus($product(3,$sum(sK3,sK0))),$sum($uminus($sum(sK0,sK2)),$sum($product(2,$sum(1,sK1)),$product($sum(1,sK1),$sum(sK0,sK2))))),1) | 0 != $sum(1,$sum($uminus($sum(sK0,sK2)),$sum($product(3,$sum(1,sK1)),$product($sum(1,sK1),$product(3,$sum(1,sK1)))))) [forward demodulation 1704,14]
1706. 0 != $sum($sum($uminus($product(3,$sum(sK0,sK3))),$sum($uminus($sum(sK0,sK2)),$sum($product(2,$sum(1,sK1)),$product($sum(1,sK1),$sum(sK0,sK2))))),1) | 0 != $sum(1,$sum($uminus($sum(sK0,sK2)),$sum($product(3,$sum(1,sK1)),$product($sum(1,sK1),$product(3,$sum(1,sK1)))))) [forward demodulation 1705,3]
1755. 7 <=> 0 = $sum(1,$sum($uminus($sum(sK0,sK2)),$sum($product(3,$sum(1,sK1)),$product($sum(1,sK1),$product(3,$sum(1,sK1)))))) [avatar definition]
1759. 8 <=> 0 = $sum($sum($uminus($product(3,$sum(sK0,sK3))),$sum($uminus($sum(sK0,sK2)),$sum($product(2,$sum(1,sK1)),$product($sum(1,sK1),$sum(sK0,sK2))))),1) [avatar definition]
1762. ~7 | ~8 [avatar split clause 1706,1759,1755]
1763. $false [avatar sat refutation 458,555,1031,1236,1425,1586,1762]
% SZS output end Proof for nia_cubes_simple_v2
% ------------------------------
% Version: Vampire 4.5.1 (commit 2f3063097 on 2021-04-08 14:14:55 +0200)
% Linked with Z3 4.8.10.0 517d907567f4283ad8b48ff9c2a3f6dce838569e z3-4.8.4-4541-g517d90756
% Termination reason: Refutation

% Memory used [KB]: 8187
% Time elapsed: 0.060 s
% ------------------------------
% ------------------------------
% Success in time 0.064 s
